# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    a=[]
    b=[]
    for char in x:
        a.append(char)
    for char in y:
        b.append(char)
    a.sort()
    b.sort()
    if a==b:
        return True
    else:
        return False
    print (a)
    print (b)
same_digits('1234','4321')
same_digits('1234','432')
```

Output:
```
False
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
for d in range(0,10):
    for c in range(0,10):
        if same_digits(d,c)==True:
            if same_digits(d,2*c)==True:
                if same_digits(d,3*c)==True:
                    if same_digits(d,4*c)==True:
                        if same_digits(d,5*c)==True:
                            if same_digits(d,6*c)==True:#i may crash everything but it was thinking that this would check all the numbers
                                return d,c
                            else:
                                return none
```

Output:
```
File "<ipython-input-33-11b6852ded2c>", line 9
    return d,c
    ^
SyntaxError: 'return' outside function
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```

```python

```